# Homework 2


## Problem 1 (Dictionary and String Methods):

You have a list of strings which contain student record. Each string contains the student's name, and their grades in Maths, Physics and Chemistry. The grades can be **floating values**. These values (names and grades) are separated by comma. You are required to save the record in a dictionary data type. The user then enters a student's name. Output the average percentage marks obtained by that student, correct to two decimal places.

Example:

    input:
        grades = ["Jake, 99, 70, 50", "Dennis, 100, 100, 98"]
        name = "Dennis"
    output: "The average grade of Dennis is 99.33"

<!-- grade_id: cell-562beecb93f03f23, points: 0, grade: false, solution: true -->
```python
def problem2(grades, name):
    output = ""
    Dict = {}
    for item in grades:
        item_new = item.split(', ')
        names = item_new[0]
        grade = item_new[1:]
        grade_int = [int(num) for num in grade]
        Dict[names] = round(sum(grade_int)/len(grade_int),2)
        
    avg_grade = Dict.get(name)
    output = "The average grade of %s is %.2f" %(name,avg_grade)
    return(output)

```

Output:
```
'The average grade of Jake is 74.33'
```

<!-- grade_id: cell-681343a4d7e22bbd, points: 15, grade: true, solution: false -->
```python
grades = ["Jake, 99, 70, 54", "Dennis, 100, 100, 98"]
name = "Dennis"
assert problem2(grades, "Dennis") == "The average grade of Dennis is 99.33"
```

<!-- grade_id: cell-402bf89f76c4bd42, points: 15, grade: true, solution: false -->
```python
grades = ["Jake, 99, 70, 55", "Dennis, 100, 100, 98"]
name = "Jake"
assert problem2(grades, "Jake") == "The average grade of Jake is 74.67"
```

## Problem 2 (List Methods):

You are given N positive numbers. Store them in a list and find the second largest number.
(If there are two largest numbers, the second largest number must be smaller than both of them)

Example:

    input: int_list = [1,3,1,4,4,2]
    ouput: 3

<!-- grade_id: cell-90ba74290831cf3b, points: 0, grade: false, solution: true -->
```python
#Write your code here:
def problem3(int_list):
    second_largest = 0
    int_list.sort(reverse=True)
    
    largest = int_list[0]
    for i in range(len(int_list)):
        if int_list[i] < largest:
            second_largest = int_list[i]
            #print(second_largest)
            break
    return(second_largest)
```

<!-- grade_id: cell-f14a368424f5d8fc, points: 15, grade: true, solution: false -->
```python
assert problem3([1,3,1,4,4,2]) == 3
```

<!-- grade_id: cell-b28263287d112648, points: 15, grade: true, solution: false -->
```python
assert problem3([1,2,3]) == 2
```

## Problem 3 (All-in-one: list + string + file read):

We downloaded several stock's daily level data for you in ./data/hw2_data/stock_data/. You will write code that takes stock tickers as inputs and output one dictinoary. The dictionary uses the stock tickers as keys and the values are these stocks' average open prices in 2017 (until the end of the file).

Exmaple:

    input: tickers = ["AAPL"]
    output: {'AAPL': 143.465799260355}


<!-- grade_id: cell-25ac6e4633eae166, points: 0, grade: false, solution: true -->
```python
def problem4(tickers):
    import os
    file_location = "./data/stock_data/"
    files = os.listdir(file_location)
    all_dict = {}
    return_dict = {}
    
    for file in files:
        ticker = file.strip(".csv")
        stock_info = open("./data/stock_data/"+file)
        lines=stock_info.readlines()
        opens = []
        for line in lines[1:]:
            daily_info = line.strip("\n").split(",")
            daily_open = eval(daily_info[2])
            opens.append(daily_open)
        avg_opens = sum(opens)/len(opens)
        all_dict[ticker] = avg_opens
    for ticker in tickers:
        return_dict[ticker] = all_dict[ticker]  
        
    return(return_dict)


```

<!-- grade_id: cell-f824aae0b179bccb, points: 20, grade: true, solution: false -->
```python
assert problem4(["AAPL"]) == {'AAPL': 143.465799260355}
```

<!-- grade_id: cell-49c8a58cfdfc3878, points: 20, grade: true, solution: false -->
```python
assert problem4(["AAPL", "GOOG"]) == {'AAPL': 143.465799260355, 'GOOG': 884.3826611893493}
```