```python
#Run this cell for your imports
import pandas as pd
import numpy as np
```

In this homework, you will writing two functions that will assign a final letter grade to each student.  The scores of each student on each assignment are given in "Final_Grades.csv". In this semester, there 6 homeworks (columns labeled HW_1, .., HW_6), 3 participation scores (column Part_1, Part_2, Part_3) and a midterm and final.

Your first task is to write a function callled "Get_Final_Perc" that has the following three inputs:

- HW_weight: The weight towards the final grade of the score on all homeworks.
- Exam_weight: The weight towards the final grade of the midterm and final. You may exam the midterm and final are weighted equally
- Part_weight: The weight towards the final grade of the score on all participation scores.

IMPORTANT:  For the HW and participation categories, the final percentage is (# points earned)/(# of points possible).  So, for example, if there were two HWs and you got 10/10 and 7/8, your final percentage on HW would be 17/18.

Inside of this function, you should

- Read in the data set into a dataframe called df_grades.
- Add to df_grades a column labeled "Final_perc", which gives the final percentage (this should be between 0 and 100 for each student) achieved by each student given the specified weighting scheme that is given as input.  Note that it may be helpful to add other columns along the way, this is okay.


A few other things that you must keep in mind when writing this function

- You should only consider students who have a score for all assignments (you can delete the rows corresponding to students with missing scores for any assignment). Hint: Recall that blanks in a csv are automatically read in as NAs.
- You may assume that for each assignment, there was at least one student who received 100% on the assignment.  Furthermore, there was no extra credit on any assignment. Hint: this is how you can figure out the max score on each assignment.

 


```python
""""
df_grades=pd.read_csv("Final_Grades.csv")

df_grades.dropna(inplace=True)
print(df_grades.head())
hw_1=df_grades["HW_1"].max()
hw_2=df_grades["HW_2"].max()
hw_3=df_grades["HW_3"].max()
hw_4=df_grades["HW_4"].max()
hw_5=df_grades["HW_5"].max()
hw_6=df_grades["HW_6"].max()
final=df_grades["Final"].max()
midterm=df_grades["Midterm"].max()
part_1=df_grades["Part_1"].max()
part_2=df_grades["Part_2"].max()
part_3=df_grades["Part_3"].max()

 
hw=hw_1+hw_2+hw_3+hw_4+hw_5+hw_6
part=part_1+part_2+part_3

 

HW_weight=40
Exam_weight=25
Part_weight=10

def hw_part(info):
    student=info[0]
    hw_total=sum(info[["HW_1","HW_2","HW_3","HW_4","HW_5","HW_6"]])
    part_total=sum(info[["Part_1","Part_2","Part_3"]])
    
    hw_per=hw_total/hw
    part_per=part_total/part
    final_per=info["Final"]/final
    midterm_per=info["Midterm"]/midterm
    
    total=hw_per*HW_weight+part_per*Part_weight+(final_per+midterm_per)*Exam_weight
    return total



df_grades["Final_perc"]=df_grades.apply(hw_part,axis=1)

df_grades
df_grades.Final_perc.mean()
""""

```

Output:
```
Student  HW_5  HW_6  Final  HW_1  Part_1  HW_2  Part_2  HW_3  Part_3  \
0  Student 1    29  28.0   43.5  18.0     1.0  29.0     1.0  18.0     1.0   
1  Student 2    30  30.0   49.5  20.0     0.0  30.0     1.0  19.0     1.0   
2  Student 3    30  30.0   45.5  20.0     1.0  30.0     1.0  20.0     1.0   
3  Student 4    28  30.0   41.0  20.0     1.0  29.0     1.0  19.0     1.0   
4  Student 5    30  30.0   46.5  20.0     1.0  30.0     1.0  20.0     1.0   

   Midterm  HW_4  
0     33.0  15.0  
1     39.0  14.0  
2     37.5  15.0  
3     33.5  14.0  
4     39.0  16.0  
88.71127425606514
```

<!-- grade_id: Get_Final_Perc, points: 0, grade: false, solution: true -->
```python
def Get_Final_Perc(HW_weight, Exam_weight, Part_weight):
    
    df_grades=pd.read_csv("Final_Grades.csv")

    df_grades.dropna(inplace=True)

    hw_1=df_grades["HW_1"].max()
    hw_2=df_grades["HW_2"].max()
    hw_3=df_grades["HW_3"].max()
    hw_4=df_grades["HW_4"].max()
    hw_5=df_grades["HW_5"].max()
    hw_6=df_grades["HW_6"].max()
    final=df_grades["Final"].max()
    midterm=df_grades["Midterm"].max()
    part_1=df_grades["Part_1"].max()
    part_2=df_grades["Part_2"].max()
    part_3=df_grades["Part_3"].max()


    hw=hw_1+hw_2+hw_3+hw_4+hw_5+hw_6
    part=part_1+part_2+part_3




    def hw_part(info):
        student=info[0]
        hw_total=sum(info[["HW_1","HW_2","HW_3","HW_4","HW_5","HW_6"]])
        part_total=sum(info[["Part_1","Part_2","Part_3"]])

        hw_per=hw_total/hw
        part_per=part_total/part
        final_per=info["Final"]/final
        midterm_per=info["Midterm"]/midterm

        total=hw_per*HW_weight+part_per*Part_weight+(final_per+midterm_per)*Exam_weight
        return total



    df_grades["Final_perc"]=df_grades.apply(hw_part,axis=1)

    
    return df_grades
```

<!-- grade_id: Test_1_1, points: 1, grade: true, solution: false -->
```python
assert np.isclose(Get_Final_Perc(40, 25, 10).Final_perc.mean(),88.711)
```

<!-- grade_id: Test_1_2, points: 1, grade: true, solution: false -->
```python
assert np.isclose(Get_Final_Perc(100, 0, 0).Final_perc.mean(),97.426)
```

<!-- grade_id: Test_1_3, points: 1, grade: true, solution: false -->
```python
assert np.isclose(Get_Final_Perc(1, 49.5, 1).Final_perc.mean(),82.310)
```

Your second and final task is to write a function called "Assign_Grades", which as the following two inputs:

- df_grades_w_final_perc: This is the df_grades dataframe with the Final_Perc column added to it.  In other words, you may assume that df_grades_w_final_perc has a column called Final_perc which is each student's final percentage.

- cutoffs: This is a list that always as 4 entries.  The first entry is the minimum score needed for an "A", the second entry is the minumum score needed for a "B", the third entry is the minumum score needed for a "C" and the last entry is the miminum score need for a "D". Any score below this fourth entry is awarded an "F".  You may assume that all cutoffs are unique so you don't have to worry about ties. Finally, you may assume that the cutoffs for better grades are always higher than the cutoffs for lower grades.

You goal is to add to df_grades_w_final_perc a column called "letter", which contains each student's letter grade (as a string, i.e. an A should added as "A") according to the inputted cutoff points.

Hint: The pandas method "cut" could be very helpful here.

```python
""""
cutoffs = [100,90,80,70,60,0]
cutoffs.sort()

score_cut=pd.cut(df_grades["Final_perc"],cutoffs,labels=["F","D","C","B","A"])

df_grades["letters"]=score_cut


df_grades
""""
```

Output:
```
Student  HW_5  HW_6  Final  HW_1  Part_1  HW_2  Part_2  HW_3  Part_3  \
0    Student 1    29  28.0   43.5  18.0     1.0  29.0     1.0  18.0     1.0   
1    Student 2    30  30.0   49.5  20.0     0.0  30.0     1.0  19.0     1.0   
2    Student 3    30  30.0   45.5  20.0     1.0  30.0     1.0  20.0     1.0   
3    Student 4    28  30.0   41.0  20.0     1.0  29.0     1.0  19.0     1.0   
4    Student 5    30  30.0   46.5  20.0     1.0  30.0     1.0  20.0     1.0   
..         ...   ...   ...    ...   ...     ...   ...     ...   ...     ...   
92  Student 93    30  30.0   52.0  20.0     1.0  30.0     1.0  20.0     1.0   
93  Student 94    30  30.0   43.0  20.0     1.0  30.0     1.0  20.0     1.0   
94  Student 95    29  30.0   47.0  20.0     1.0  30.0     1.0  20.0     1.0   
95  Student 96    30  30.0   47.0  20.0     1.0  30.0     1.0  20.0     1.0   
96  Student 97    30  28.0   40.0  18.0     1.0  29.0     1.0  20.0     1.0   

    Midterm  HW_4  Final_perc letters letter  
0      33.0
```

<!-- grade_id: Assign_Grades, points: 0, grade: false, solution: true -->
```python
def Assign_Grades(df_grades_w_final_perc, cutoffs):
    
    cutoffs.append(101)
    cutoffs.append(0)
    cutoffs.sort()

    score_cut=pd.cut(df_grades_w_final_perc["Final_perc"],cutoffs,labels=["F","D","C","B","A"], right=False)
    df_grades_w_final_perc["letter"]=score_cut

    
    return df_grades_w_final_perc.letter.value_counts()
```

<!-- grade_id: Test_2_1, points: 1, grade: true, solution: false -->
```python
df_grades = Get_Final_Perc(40, 25, 10)
cutoffs = [90,80,70,60]
assert Assign_Grades(df_grades, cutoffs)["A"]==49
```

<!-- grade_id: Test_2_2, points: 1, grade: true, solution: false -->
```python
df_grades = Get_Final_Perc(100, 0, 0)
cutoffs = [100,80,70,60]
assert Assign_Grades(df_grades, cutoffs)["A"]==27
```