# Midterm: Recommender System for Movies

### (Note: This midterm assignment will have hidden test cases)

In this project, you will implement a recommender system for your classmates, professor and TAs based on the movie survey we have conducted. The movie preference file is at **./data/movie_preference.csv**

## Recommender System

The objective of a Recommender System is to recommend relevant items for users, based on their preference. Recommender system is prevalent in the digital space. For example, when you go shopping on Amazon, you will notice that Amazon is recommending products on the front page before you even type anything in the search box. Similarly, when you go on YouTube, the top bar of Youtube is typically "videos recommended to you." All these features are based on recommmender system. 

What item to recommend to which user is arguably the most important business decision in many digital platforms. For instance, YouTube cannot control which videos that users upload to it. It cannot control which videos users like to watch. Moreoveor, since watching videos is free, YouTube cannot change the price of its items. It does not have inventory either since each video can be viewed as many times as possible. In this case, what could YouTube control? Or in other words, what differentiates a good video streaming service from a bad one? The answer is recommender system. 

## Types of Recommender Systems

There are **three** types of recommender system. **In this bonus project, we will implement the first one.**

### Popularity-based Recommendation 

The most obvious system is popularity-based recommendation. In this case, this model recommends to a user the most popular items that the user has not previously consumed. In the movie setting, we will recommend the movie that most users have liked and consumed. In other words, this system utilizes the "widom of the crowds." It usually provides good recommendations for most of the people. Since it is easy to implement, people normally use popularity-based recommendation as a baseline. *Note: this system is not personalized. If both consumers did not watch Movie A and Movie A is the most popular one, both of them will be recommended Movie A.*

### Content-based Recommendation 

This recommender system leverages the data of one customer's historical actions. This recommender systems first utilizes a set of features to describe an item (for example, for movies, we can use the movie's director, main actor, main actress, genre, etc. to describe the movie). When a user comes in, the system will recommend the movies that are closest to the movie that the users have consumed and liked before in terms of the features. For instance, if a user likes action move from Nolan the most, this system will recommend another action movie from Nolan that this user has not consumed. *Note: we will not implement this system in this bonus project since it requires knowledge about supervised learning. We will come back to this topic at the end of this semester.*

### Collaborative Filtering Recommendation

The last type of recommender system is called collaborative filtering. This approach uses the memory of previous users interactions to compute users similarities based on items they've interacted (user-based approach) or compute items similarities based on the users that have interacted with them (item-based approach).

A typical example of this approach is User Neighbourhood-based CF, in which the top-N similar users (usually computed using Pearson correlation) for a user are selected and used to recommend items those similar users liked, but the current user have not interacted yet. 

## Step-0 Read-in the preference file

The first exercise is to read in the movie preference csv file (you need to use relative path). 

It returns two things:

1. A dictionary where the key is username and the value is a vector of (-1, 0, 1) that indicates the users preference across movies (in the order of the csv file). 

2. A list of strings that contains movie names. (The order of movie names should be the same as the order in the original csv file)


**Note 1:** Your result should exactly match the results from the assert statements. This means you should pay attention to extra space, newline, etc.

**Note 2:** If there are two records with the same name, use the first record from the person.


<!-- grade_id: cell-b3679d3e7a94fc65, points: 0, grade: false, solution: true -->
```python
def read_in_movie_preference():
    """Read the move data, and return a 
    preference dictionary."""
    preference = {}
    movies = []
    
    path="./data/preference.csv"
    file=open(path,"r")
    lines=file.readlines()
    titles=lines[0].strip("\n").split(",")
    
    for string in titles[2:]:
        movies.append(string.strip(" "))
    
    n_movies=len(movies)
    
    for line in lines[1:]:
        line=line.strip("\n").split(",")
        username=line[1]
        scores=[]
        for i in range(n_movies):
            score=int(line[i+2])
            scores.append(score)
        if username not in preference.keys():
            preference[username]=scores
        
  
    return [movies, preference]

```

```python
#draft
preference = {}
movies = []

path="./data/preference.csv"
file=open(path,"r")
lines=file.readlines()
titles=lines[0].replace("\n","").split(",")

for string in titles[2:]:
    movies.append(string.strip(" "))
print(movies)    
#info_1=lines[1].strip("\n").split(",")
#username_1=info_1[1]
#username_1
#prefer=info_1[2:]
#prefer
n_movies=len(movies)

for line in lines[1:]:
    line=line.replace("\n","").split(",")
    username=line[1]
    scores=[]
    for i in range(n_movies):
        score=int(line[i+2])
        scores.append(score)
    preference[username]=scores
preference["Ziqing Ouyang"]
preference["Jacob Scheinman"]
x=['The Shawshank Redemption', 'The Godfather',
                       'The Dark Knight', 'Star Wars: The Force Awakens',
                       'The Lord of the Rings: The Return of the King',
                       'Inception', 'The Matrix', 'Avengers: Infinity War',
                       'Interstellar', 'Spirited Away', 'Coco', 'The Dark Knight Rises',
                       'Braveheart', 'The Wolf of Wall Street', 'Gone Girl', 'La La Land',
                       'Shutter Island', 'Ex Machina', 'The Martian', 'Kingsman: The Secret Service']
print(movies==x)
```

Output:
```
['The Shawshank Redemption', 'The Godfather', 'The Dark Knight', 'Star Wars: The Force Awakens', 'The Lord of the Rings: The Return of the King', 'Inception', 'The Matrix', 'Avengers: Infinity War', 'Interstellar', 'Spirited Away', 'Coco', 'The Dark Knight Rises', 'Braveheart', 'The Wolf of Wall Street', 'Gone Girl', 'La La Land', 'Shutter Island', 'Ex Machina', 'The Martian', 'Kingsman: The Secret Service']
True
```

<!-- grade_id: cell-186a478565fdc790, points: 5, grade: true, solution: false -->
```python
[movies, preference] = read_in_movie_preference()
assert len(movies) == 20
```

<!-- grade_id: cell-5630d404a280c24d, points: 5, grade: true, solution: false -->
```python
[movies, preference] = read_in_movie_preference()
assert movies == ['The Shawshank Redemption', 'The Godfather',
                       'The Dark Knight', 'Star Wars: The Force Awakens',
                       'The Lord of the Rings: The Return of the King',
                       'Inception', 'The Matrix', 'Avengers: Infinity War',
                       'Interstellar', 'Spirited Away', 'Coco', 'The Dark Knight Rises',
                       'Braveheart', 'The Wolf of Wall Street', 'Gone Girl', 'La La Land',
                       'Shutter Island', 'Ex Machina', 'The Martian', 'Kingsman: The Secret Service']
```

<!-- grade_id: cell-b0515b35a5476ae0, points: 10, grade: true, solution: false -->
```python
[movies, preference] = read_in_movie_preference()
assert preference["Jacob Scheinman"] == [1, 1, 1, 1, 1, 1, 1, 1, -1, 0, 1, 1, 0, 1, 0, 0, 0, 0, 1, 1]
assert preference["Ziqing Ouyang"] == [0, 0, 0, 0, 0, 1, 1, 1, 0, 0, 0, 0, 0, 0, 0, 1, 0, 0, 0, 1]
```

## Step-1 Popularity-based Ranking
### Step-1.1 Compute the ranking of most popular movies

Your next task is to take the movie preference dataframe and computes the popular ranking of movies from the most popular to the least popular. You should return a list where each element represents the popularity ranking of the movies. The order of the list should reflect the order of the movie names in the dataframe. 

In the process to compute a movie's popularity. You should first compute how many times people have liked movies in the entire dataset across all movies. You should then compute how many times people have disliked movies in the entire data set across all movies.

Assuming that people have liked movies **A** times in the entire data set and disliked movies **B** times in the entire data set. The popularity of a movie is then defined as **Num_of_People_Like_the_Movie - A / B * Num_of_People_Dislike_the_Movie**

(We use A/B to normalize the weights of likes and dislikes because if one type of reaction is rare, it derseves more weights. For example, if all movies on average are liked 1000 times but disliked only once. Then the signal of dislike on a movie's quality should be much stronger than the signal of likes on a movie's quality).


Your function should return:
    1. A dictionary where the key are movie names and the values are correpsonding movie popularity.
    2. A list of movie names sorted ascendingly by their popularity. For example, if 'The Shawshank Redemption' is the second most popular movie, the second element in the list should be 'The Shawshank Redemption'.
    3. **A** and **B** as defined above.

**Note: You may want to use prior functions to help you read data inside this function**

<!-- grade_id: cell-330ba96678776e4a, points: 0, grade: false, solution: true -->
```python
def movies_popularity_ranking():
    movie_popularity = {}
    movie_popularity_rank = []
    total_likes = 0
    total_dislikes = 0
    like_list=[]
    dislike_list=[]    

    for i in range(len(movies)):
        likes=0
        dislikes=0
        for key in preference.keys():
            info=preference[key]
            if info[i]==1:
                likes=likes+1
            elif info[i]==-1:
                dislikes=dislikes+1
        like_list.append(likes)
        dislike_list.append(dislikes)

    total_likes=sum(like_list)
    total_dislikes=sum(dislike_list)
    A=total_likes
    B=total_dislikes

    for i in range(len(movies)):
        likes=0
        dislikes=0
        for key in preference.keys():
            info=preference[key]
            if info[i]==1:
                likes=likes+1
            elif info[i]==-1:
                dislikes=dislikes+1
        popularity=likes-(A/B)*dislikes
        movie_popularity[movies[i]]=popularity

        
    movies_sorted=sorted (movie_popularity.items(), key=lambda item:item[1], reverse = True )

    for i in range(len(movies_sorted)):
        movie_info=movies_sorted[i]
        movie_popularity_rank.append(movie_info[0])
    
    
    return movie_popularity, movie_popularity_rank, total_likes, total_dislikes
```

```python
#draft
movie_popularity = {}
movie_popularity_rank = []
like_list=[]
dislike_list=[]
total_likes = 0
total_dislikes = 0

for i in range(len(movies)):
    likes=0
    dislikes=0
    for key in preference.keys():
        info=preference[key]
        if info[i]==1:
            likes=likes+1
        elif info[i]==-1:
            dislikes=dislikes+1
    like_list.append(likes)
    dislike_list.append(dislikes)
    
total_likes=sum(like_list)
total_dislikes=sum(dislike_list)

for i in range(len(movies)):
    likes=0
    dislikes=0
    for key in preference.keys():
        info=preference[key]
        if info[i]==1:
            likes=likes+1
        elif info[i]==-1:
            dislikes=dislikes+1
    popularity=likes-(A/B)*dislikes
    movie_popularity[movies[i]]=popularity

print(A,B)
print(movie_popularity)

movies_sorted=sorted (movie_popularity.items(), key=lambda item:item[1], reverse = True )
print(movies_sorted)
for i in range(len(movies_sorted)):
    movie_info=movies_sorted[i]
    movie_popularity_rank.append(movie_info[0])
    
movie_popularity_rank
```

Output:
```
1300 236
{'The Shawshank Redemption': 66.98305084745763, 'The Godfather': -8.593220338983059, 'The Dark Knight': 4.881355932203377, 'Star Wars: The Force Awakens': -90.22033898305085, 'The Lord of the Rings: The Return of the King': -40.152542372881356, 'Inception': 59.45762711864407, 'The Matrix': 24.44067796610169, 'Avengers: Infinity War': 14.86440677966101, 'Interstellar': 7.389830508474574, 'Spirited Away': -1.6101694915254257, 'Coco': 22.423728813559322, 'The Dark Knight Rises': 10.898305084745758, 'Braveheart': -13.067796610169495, 'The Wolf of Wall Street': 27.40677966101694, 'Gone Girl': -31.101694915254242, 'La La Land': -15.118644067796623, 'Shutter Island': -21.59322033898306, 'Ex Machina': -48.610169491525426, 'The Martian': 1.8983050847457577, 'Kingsman: The Secret Service': 29.423728813559322}
[('The Shawshank Redemption', 66.98305084745763), ('Inception', 59.45762711864407), ('Kingsman: The Secret Service', 29.423728813559322), ('The Wolf of Wall Street', 27.40677966101
```

<!-- grade_id: cell-3f47ab3b02050bd4, points: 5, grade: true, solution: false -->
```python
movie_popularity, movie_popularity_rank, total_likes, total_dislikes = movies_popularity_ranking()
assert total_likes == 1300
assert total_dislikes == 236
```

<!-- grade_id: cell-c306578d650b19c0, points: 5, grade: true, solution: false -->
```python
movie_popularity, movie_popularity_rank, total_likes, total_dislikes = movies_popularity_ranking()
assert round(movie_popularity["The Shawshank Redemption"], 2) == 66.98
assert round(movie_popularity["Avengers: Infinity War"], 2) == 14.86
```

<!-- grade_id: cell-736c1f24d6b1e609, points: 5, grade: true, solution: false -->
```python
movie_popularity, movie_popularity_rank, total_likes, total_dislikes = movies_popularity_ranking()
assert movie_popularity_rank == ['The Shawshank Redemption',
 'Inception',
 'Kingsman: The Secret Service',
 'The Wolf of Wall Street',
 'The Matrix',
 'Coco',
 'Avengers: Infinity War',
 'The Dark Knight Rises',
 'Interstellar',
 'The Dark Knight',
 'The Martian',
 'Spirited Away',
 'The Godfather',
 'Braveheart',
 'La La Land',
 'Shutter Island',
 'Gone Girl',
 'The Lord of the Rings: The Return of the King',
 'Ex Machina',
 'Star Wars: The Force Awakens']
```

## 1.2 Recommendation

You will then implement a recommendation function. This function will take in a user's name, it will return a string representing the name of the top movie that this user has not watched and has best popularity ranking (i.e., lowest ranking number) only if this unwatched movie has higher popularity scores than the average of popularity scores of movies that this user has watched (regardless whether he/she likes or dislikes the movie). 

If the user name does not exit, this function should return "Invalid user."

If the user has watched all movies, this function should return "Unfortunately, no new movies for you."

If the unwatched movies all have lower popularity scores than the average score of movies watched by this user, this function should return "Unfortunately, no new movies for you."

**Note: Again, you may want to use prior functions to help you read data and rank movies inside this function**

<!-- grade_id: cell-3cf2c39ef6f2c170, points: 0, grade: false, solution: true -->
```python
def Recommendation(name):
    
    recommended_movie = ""
    if name in preference.keys():
        info=preference[name]
        #print(info)
        #calculate the average of popularity scores of movies that this user has watched 
        watched_movies=[]
        watched_avg=0

        for i in range(len(info)):
            if info[i]!=0:
                watched_movies.append(movie_popularity[movies[i]])
            if len(watched_movies)!=0:
                watched_avg=sum(watched_movies)/len(watched_movies)

        #print(watched_avg)

        #order the info in accordance with movie_popularity_rank
        dict_info={}
        for i in range(len(info)):
            dict_info[movies[i]]=info[i]

        #print(dict_info)

        #compare the avg with the unwatched movies
        if len(watched_movies)==len(movies):#watched all
            return "Unfortunately, no new movies for you."
        elif len(watched_movies)<len(movies) and len(watched_movies)!=0:#watched some
            for target_movie in movie_popularity_rank:
                if dict_info[target_movie]==0 and movie_popularity[target_movie]>=watched_avg:
                    recommended_movie=target_movie
                    break
            if recommended_movie=="":#watched some but all < avg
                return "Unfortunately, no new movies for you."
        elif len(watched_movies)==0:#watched none
            recommended_movie=movies[0]
    else:
        return "Invalid user."
    #print(recommended_movie)
    
    return recommended_movie

```

```python
#是草稿
recommended_movie = ""
name="Test Student 2"
if name in preference.keys():
    info=preference[name]
    #print(info)
    #calculate the average of popularity scores of movies that this user has watched 
    watched_movies=[]

    for i in range(len(info)):
        if info[i]!=0:
            watched_movies.append(movie_popularity[movies[i]])
        if len(watched_movies)!=0:
            watched_avg=sum(watched_movies)/len(watched_movies)

    #print(watched_avg)

    #order the info in accordance with movie_popularity_rank
    dict_info={}
    for i in range(len(info)):
        dict_info[movies[i]]=info[i]

    #print(dict_info)

    #compare the avg with the unwatched movies
    if len(watched_movies)==len(movies):#watched all
        print("Unfortunately, no new movies for you.")
    elif len(watched_movies)<len(movies) and len(watched_movies)!=0:#watched some
        for target_movie in movie_popularity_rank:
            if dict_info[target_movie]==0 and movie_popularity[target_movie]>=watched_avg:
                recommended_movie=target_movie
                break
        if recommended_movie=="":#watched some but all < avg
            print("Unfortunately, no new movies for you.")   
    elif len(watched_movies)==0:#watched none
        recommended_movie=movies[0]
else:
    print("Invalid user.")
#print(recommended_movie)
Recommendation("Test Student 2") == 'Invalid user.'
```

Output:
```
Invalid user.
True
```

<!-- grade_id: cell-13b32edf438b5c7e, points: 5, grade: true, solution: false -->
```python
assert Recommendation("Jiaxu Rong") == 'Inception'
assert Recommendation("Nobody") == 'The Shawshank Redemption'
```

<!-- grade_id: cell-17b895dde51f52ea, points: 5, grade: true, solution: false -->
```python
assert Recommendation("Dennis Zhang") == 'Kingsman: The Secret Service'
```

<!-- grade_id: cell-c33a47a80a692472, points: 5, grade: true, solution: false -->
```python
assert Recommendation("Test Student 2") == 'Invalid user.'
```

## 2.1 Cosine Similarity
Let us then use collaborative filtering to find the recommendation.

First, we need to get the cosine similarity beween movies and users. Again, we can use the preference file that we get in Step 0. In that case, each person is represented by a vector of (0, 1, -1). Cosine similarity in our case is the dot product of the two preference vectors divided by the product of the magnitude of the two preference vectors. In other words, if person A has preference vector A, and person B has preference vector B, their cosine similarity is equal to 

$$ \frac{A \cdot B}{||A||||B||} = \frac{\sum_i^n A_iB_i}{\sqrt{\sum_i^nA_i^2}\sqrt{\sum_i^nB_i^2}}$$

If a person has not watched any movies, then the cosine similarity between this person and any other person is defined as 0. For more information on cosine simialrity, you can read this [wiki page](https://en.wikipedia.org/wiki/Cosine_similarity)

For example, the following two vectors represent Dennis' and Jake's preference over 3 movies.

         Inception  Coco     The Dark Knight
    Jake     1         -1        0

    Dennis   -1         0        1

In this case, Dennis and Jake's cosine similarity is equal to

$$ \frac{1*(-1)+(-1)*0+0*(-1)}{\sqrt{1+1}*\sqrt{(-1)^2+1}} = \frac{-1}{2} = -0.5$$

**Your task** is to write a similarity function that takes in two names and returns the jaccard similarity between these two people. If one or two names do not exist in the database, return 0. 



<!-- grade_id: cell-f32c4bc84d3b8cb5, points: 0, grade: false, solution: true -->
```python
def Similarity(name_1, name_2):
    """Given two names and preference, get the similarity 
    between two people"""
    cosine = 0
    if name_1 in preference.keys() and name_2 in preference.keys():    
    
        info_1=preference[name_1]
        info_2=preference[name_2]

        numerator=0
        magnitude_1=0
        magnitude_2=0
        for i in range(len(info_1)):
            numerator=numerator+info_1[i]*info_2[i]
            magnitude_1=magnitude_1+info_1[i]**2
            magnitude_2=magnitude_2+info_2[i]**2
        if magnitude_1!=0 and magnitude_2!=0:
            cosine=numerator/((magnitude_1**0.5)*(magnitude_2**0.5))    
        
        
    
    return cosine
```

```python
#test
name_1="Test Student"
name_2="Nobody"

info_1=preference[name_1]
info_2=preference[name_2]

numerator=0
magnitude_1=0
magnitude_2=0
for i in range(len(info_1)):
    numerator=numerator+info_1[i]*info_2[i]
    magnitude_1=magnitude_1+info_1[i]**2
    magnitude_2=magnitude_2+info_2[i]**2   
cosine=numerator/((magnitude_1**0.5)*(magnitude_2**0.5))
print(magnitude_1,magnitude_2)
cosine
```

Output:
```
7 5
0.16903085094570328
```

<!-- grade_id: cell-7a65b16620e5bb27, points: 10, grade: true, solution: false -->
```python
assert round(Similarity("Test Student", "Nobody"), 2) == 0.17
assert round(Similarity("Test Student", "DJZ2"), 2) == -0.27
```

<!-- grade_id: cell-95386d67a475aa08, points: 5, grade: true, solution: false -->
```python
assert round(Similarity("Test Student", "Test Student 2"), 2) == 0
```

## 2.2 Movie Soulmate

Your next task is to find the movie soulmate for a person. In order to find a person's movie soulmate, you will compute the cosine similarity between this person and every other person in the data set. You will then return the person who has the highest cosine similarity with the focal person. If two people have the same cosine similarity with the focal person, you can tie break by the length of names (the name with lower length will be the soulmate). If the focal person does not exist in the database, return an empty string as the soulmate name.

You function will return two things:
1. the name of the soulmate
2. the largest cosine similarity



<!-- grade_id: cell-c53815c71ec1252d, points: 0, grade: false, solution: true -->
```python
def Movie_Soul_Mate(name):
    """Given a name, get the player that has highest Jaccard 
    similarity with this person."""
    soulmate = ""
    cosine_similarity = -100
    
    for person in preference.keys():
        if person != name:
            if Similarity(name, person)>cosine_similarity:
                cosine_similarity=Similarity(name, person)
                soulmate=person
            elif Similarity(name, person)==cosine_similarity:
                if len(person)<len(soulmate):
                    soulmate=person

    return soulmate, cosine_similarity

```

```python
#test
name="Test Student"
soulmate = ""
cosine_similarity = -100

for person in preference.keys():
    if person != name:
        if Similarity(name, person)>cosine_similarity:
            cosine_similarity=Similarity(name, person)
            soulmate=person
        elif Similarity(name, person)==cosine_similarity:
            if len(person)<len(soulmate):
                soulmate=person
        
print(cosine_similarity,soulmate)        
    
    
```

Output:
```
0.8017837257372731 Michael Treiber
```

<!-- grade_id: cell-316ca2c5633dd4c5, points: 5, grade: true, solution: false -->
```python
soulmate, cosine_similarity = Movie_Soul_Mate("Q")
assert soulmate == 'Yunong Tian'
assert round(cosine_similarity, 2) == 0.75
```

<!-- grade_id: cell-dce0505fc8de593e, points: 10, grade: true, solution: false -->
```python
soulmate, cosine_similarity = Movie_Soul_Mate("Test Student")
assert soulmate == 'Michael Treiber'
assert round(cosine_similarity, 2) == 0.80
```

<!-- grade_id: cell-c96560dd94b9fcb0, points: 5, grade: true, solution: false -->
```python
soulmate, cosine_similarity = Movie_Soul_Mate("Yunong Tian")
assert soulmate == 'Yuchen'
assert round(cosine_similarity, 2) == 0.81
```

## 2.3 Memory-based Collaborative Filtering Recommendation

Now after finding a person's movie soulmate, we can then construct a (very preliminary) collaborative filtering recommendation. In our recommendation system, for a focal person, we first find his or her soul mate. We then find all the movies that he/she has not watched but the soul mate has watched and liked. Among all of these movies, we recommend the movie with the highest popularity ranks defined in Step 1.1 and 1.2. 


Again, 

if the user name does not exit, this function should return "Invalid user."

If the person has watched all the movies, return "Unfortunately, no new movies for you." 

If there is no movies watched and liked by the soulmate but not watched by the focal person, then return the movie (or string) that should be returned in Step 1.2.


<!-- grade_id: cell-48fc508fee5f75a5, points: 0, grade: false, solution: true -->
```python
def Recommendation2(name):
    recommended_movie = ""
    
    if name in preference.keys():
        soulmate=Movie_Soul_Mate(name)[0]
        #print(Movie_Soul_Mate(name)[0])
        info_target=preference[name]
        info_soulmate=preference[soulmate]

        dict_infos={}
        for i in range(len(info_target)):
            dict_infos[movies[i]]=[info_target[i],info_soulmate[i]]
        dict_not_watched={}
        for movie in dict_infos.keys():
            if dict_infos[movie][0]==0 and dict_infos[movie][1]!=0:
                dict_not_watched[movie]=[0,dict_infos[movie][1]]

        for movie in movie_popularity_rank:
            if movie in dict_not_watched.keys() and dict_not_watched[movie][1]==1:
                recommended_movie=movie
                break
            else:
                recommended_movie=Recommendation(name)
    else:
        return "Invalid user."
    
    return recommended_movie

```

```python
name="Test Student Long Name"
recommended_movie = ""

if name in preference.keys():
    soulmate=Movie_Soul_Mate(name)[0]
    #print(Movie_Soul_Mate(name)[0])
    info_target=preference[name]
    info_soulmate=preference[soulmate]
    
    dict_infos={}
    for i in range(len(info_target)):
        dict_infos[movies[i]]=[info_target[i],info_soulmate[i]]
    dict_not_watched={}
    for movie in dict_infos.keys():
        if dict_infos[movie][0]==0 and dict_infos[movie][1]!=0:
            dict_not_watched[movie]=[0,dict_infos[movie][1]]
    
    for movie in movie_popularity_rank:
        if movie in dict_not_watched.keys() and dict_not_watched[movie][1]==1:
            recommended_movie=movie
            break
        else:
            recommended_movie=Recommendation(name)
else:
    return "Invalid user."

```

Output:
```
Michael Treiber
The Shawshank Redemption
```

<!-- grade_id: cell-128ba93274b3f85f, points: 5, grade: true, solution: false -->
```python
assert Recommendation2("Test Student") == 'Inception'
assert Recommendation2("Test Student Long Name") == 'The Shawshank Redemption'
```

<!-- grade_id: cell-8490b270b1ff1927, points: 10, grade: true, solution: false -->
```python
assert Recommendation2("Test Student Long Name") == 'The Shawshank Redemption'
```